Evidence of social learning across symbolic cultural barriers in sperm whales

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #2 (Public review):

Summary:

The current article presents a new type of analytical approach to the sequential organisation of whale song units.

Strengths:

The detailed description of the internal temporal structure of whale songs is something that has been thus far lacking.

Weaknesses:

The conceptual and terminological bases of the paper are problematical and hamper comparison with other taxa, including humans. According to signal theory, codas are indexical rather than symbolic. They signal an individual's group identity. Borrowing from humans and linguistics, coda inter-group variation represents a case of accents - phonologically different varieties of the same call - not dialects, confirming they are an index. This raises serious doubt about whether alleged "symbolism" and similarity between whale and human vocal behaviour is factual.

We respect that the reviewer does not agree with describing codas as symbolic markers of cultural identity in sperm whales, but ultimately we find the quantitative evidence presented in Hersh et al. (2022) compelling, and stand by the framing of our manuscript, which builds on this foundation.

The same applies to the difference between ICIs (inter-click interval) and IOIs (inter-onset interval). If the two are equivalent, variation in click duration needs to be shown so small that can be considered negligible. This raises serious doubt about whether the alleged variation in whale codas is indeed rhythmic in nature and prevents future efforts for comparison with the vocal capacities of other species. The scope and relevance of this paper for the broader field is limited.

We believe there has been a miscommunication. Coda inter-click intervals are calculated as the time between the onsets of sequential clicks within a coda. This is identical to definitions of inter-onset intervals in many publications, including:

Burchardt and Knörnschild (2020): “the duration between the beginning of one element and the next”

Friberg and Battel (2002): “the time interval between the onset of the tone and the onset of the immediately following tone”

De Gregorio et al. (2021): “the time between the onset of a note and the next one”

In response to a comment from this reviewer in the first round of revisions, we made the point that we do not believe rhythm analyses need be restricted to inter-onset intervals alone. Regardless of that stance, we did analyze inter-onset intervals in this manuscript and accordingly are capturing aspects of rhythm in our analyses. We have removed a poorly worded sentence in our introduction and apologize for any confusion it caused. We have also made this explicit in lines 30–35: “This classification is based on the total number of clicks and their rhythm and tempo extrapolated from the time interval between the onsets of consecutive clicks: the inter-click interval (ICI) [15, 16] (Fig. 1A). This measure is equivalent to the inter-onset intervals (IOIs) often used in rhythm analyses [17, 18, 19] but for the sake of compatibility with studies on sperm whale acoustics, we use ICI terminology throughout this paper.”

In our analyses, inter-click intervals and inter-onset intervals are equivalent measures.

Recommendations for the authors:

Reviewer #2 (Recommendations for the authors):

My concerns regarding interdisciplinary terminology and methods remain unaddressed. The study's inaccurate terminology hinders reliable comparison with other taxa, including humans. Being "symbolic" bears no weight on the new method that the authors present, thus, the unwillingness for compatibility is limiting and perplexing. The authors state that codas have been previously described as being symbolic, but just because poor terminology has been used before doesn't justify perpetuating it, especially when it confounds and conflicts with broader comparative efforts.

We agree that being symbolic bears no weight on the new method we present, but we believe it does bear weight on our interpretation of what our method reveals about patterns in sperm whale communication. For that reason, we have opted to maintain the current framing of our manuscript.

The same applies to the difference between ICIs and IOIs. The authors resist amending terminology, even though they state the two represent the same measure. If so, want prevents the correct use of IOIs?

We have opted to use ICI throughout the paper because it is standard terminology in sperm whale acoustics, but we have now made the ICI/IOI equivalence explicitly clear in the introduction.

References:

Burchardt LS, Knörnschild M. 2020. Comparison of methods for rhythm analysis of complex animals’ acoustic signals. PLoS Computational Biology 16. doi:10.1371/journal.pcbi.1007755

De Gregorio C, Valente D, Raimondi T, Torti V, Miaretsoa L, Friard O, Giacoma C, Ravignani A, Gamba M. 2021. Categorical rhythms in a singing primate. Current Biology 31:R1379–R1380. doi:10.1016/j.cub.2021.09.032

Friberg A, Battel GU. 2002. Structural communication In: Parncutt R, McPherson G, editors. The Science & Psychology of Music Performance: Creative Strategies for Teaching and Learning. Oxford University Press. doi:10.1093/acprof:oso/9780195138108.001.0001

Hersh TA, Gero S, Rendell L, Cantor M, Weilgart L, Amano M, Dawson SM, Slooten E, Johnson CM, Kerr I, Payne R, Rogan A, Andrews O, Ferguson EL, Hom-Weaver CA, Norris TF, Barkley YM, Merkens KP, Oleson EM, Doniol-Valcroze T, Pilkington J, Gordon J, Fernandes M, Guerra M, Hickmott L, Whitehead H. 2022. Evidence from sperm whale clans of symbolic marking in non-human cultures. Proceedings of the National Academy of Sciences 119:e2201692119. doi:10.1073/pnas.2201692119

Changes from the earlier version to the revised version

Abstract

Removed in the revised version: Abstract

Earlier version: We provide quantitative evidence suggesting social learning in sperm whales across sociocultural boundaries, using acoustic data from the Pacific and Atlantic Oceans. Traditionally, sperm whale populations are categorized into clans based on their vocal repertoire: the rhythmically patterned click sequences (codas) that they use. Among these codas, identity codas function as symbolic markers for each clan, accounting for 35-60% of codas they produce. We introduce a computational method to model whale communication, which encodes rhythmic micro-variations within codas, capturing their vocal style. We find that vocal style-clans closely align with repertoire-clans. However, contrary to vocal repertoire, we show that sympatry increases vocal style similarity between clans for non-identity codas, i.e. most codas, suggesting social learning across cultural boundaries. More broadly, this subcoda structure model offers a framework for comparing communication systems in other species, with potential implications for deeper understanding of vocal and cultural transmission within animal societies.

Revised version: We provide quantitative evidence suggesting social learning in sperm whales across socio-cultural boundaries, using acoustic data from the Pacific and Atlantic Oceans. Traditionally, sperm whale populations are categorized into clans based on their vocal repertoire: the rhythmically patterned click sequences (codas) that they use. Among these codas, identity codas function as symbolic markers for each clan, accounting for 35-60% of codas they produce. We introduce a computational method to model whale communication, which encodes rhythmic micro-variations within codas, capturing their vocal style. We find that vocal style-clans closely align with repertoire-clans. However, contrary to vocal repertoire, we show that sympatry increases vocal style similarity between clans for non-identity codas, i.e. most codas, suggesting social learning across cultural boundaries. More broadly, this subcoda structure model offers a framework for comparing communication systems in other species, with potential implications for deeper understanding of vocal and cultural transmission within animal societies.

Introduction

Cultural transmission is broadly defined as the transmission of information or behaviors between individuals of the same species by means of social learning [1, 2]. While humans represent a benchmark of such capacity, cultural transmission has been observed in a wide variety of animals, including cetaceans [3, 4], songbirds [5], non-human primates [6], and insects [7]. It typically takes one of three forms: vertical transmission, from adult kin to young kin; oblique transmission, from unrelated adults to young; or horizontal transmission, from peer to peer [8].

When animals have the capacity for social learning, group-specific differences can arise and remain stable when they become distinguishable by symbolic markers: arbitrary group-identity signals that are recognizable by both members of the group itself and by members of other groups [9, 10]. In humans, symbolic markers can take a myriad of forms, ranging from visible signs, such as tattoos or garments, to communication cues or signals, such as idiomatic sentences or accents [9, 10, 11]. In animals, however, quantitative evidence of symbolic markers is remarkably scarce, one exception being recent results on the use of identity codas in sperm whale social communication [12].

Added in the revised version: Sperm whales live in multi-tiered societies and have a complex vocal communication system [14]. They communicate through rhythmic patterns (codas) of short broadband sounds (clicks), which have traditionally been classified into a set of coda types. This classification is based on the total number of clicks and their rhythm and tempo extrapolated from the time interval between the onsets of consecutive clicks: the inter-click interval (ICI) [15, 16] (Fig. 1A). This measure is equivalent to the inter-onset intervals (IOIs) often used in rhythm analyses [17, 18, 19] but for the sake of compatibility with studies on sperm whale acoustics, we use ICI terminology throughout this paper. For example, the 4-regular (4R2) type refers to a pattern of four evenly spaced clicks, whereas the 1+3 type refers to two clicks separated by a longer pause followed by two clicks in quick succession. Coda types are thus standardized rhythmic patterns, but individual vocalizations of a given coda type exhibit micro-variations around that pattern. In the case of sperm whales ICIs can be used instead of IOIs due to the very low variance in duration of the individual clicks [20, 21].

Statistical modeling of subcoda structure in sperm whales.

Earlier version: A Sperm whale communication consists of rhythmic sequences of clicks, called codas. A coda is specified by a sequence of inter-click intervals (ICIs). Codas are classified into types based on their rhythmic pattern, which can have various degrees of regularity (e.g., 4R2 vs 1+3). B Social groups of sperm whales employ specific vocal repertoires: the set of coda types they use and their associated usage frequencies. As an illustration, we show those of the EC1 and EC2 clans from the Dominica dataset [13]. Only the most numerous coda types are shown: the rest of the vocal repertoires consists of more coda types with residual frequencies. C The subcoda structure can be modeled by considering rhythmic variations within codas of the same coda type. To do so, codas are represented as sequences of discrete inter-click intervals (dICIs), by discretizing absolute ICIs into discrete bins, which can then be considered akin to symbols (e.g. A, B, C, …), providing a tokenization of codas. Different instances of a single coda type can correspond to slightly different dICI sequences. The resulting dICI sequences are modeled using variable-length Markov chains, which can be represented as subcoda trees. These trees are built starting from the root node. Each node represetns a dICI and a dICI is a child of another in the tree if it’s presence in the affects the output probability of the parent. That is the probability of the parent depends on wether the child node has been observed previously in the sequence; see the Methods for a thorough explanation. These trees can be built for an individual speaker or for a group of speakers (here represented by colors), and capture the statistical properties of rhythm variations within codas and in the transitions between those. In other words, the tree captures a vocal style—how they say what they say. The vocal styles of different groups of sperm whales can be quantitatively compared by calculating a distance between their subcoda trees, which is a statistical distance between the dICI output distributions.

Revised version: A Sperm whale communication consists of rhythmic sequences of clicks, called codas. A coda is specified by a sequence of inter-click intervals (ICIs). Codas are classified into types based on their rhythmic pattern, which can have various degrees of regularity (e.g., 4R2 vs 1+3). B Social groups of sperm whales employ specific vocal repertoires: the set of coda types they use and their associated usage frequencies. As an illustration, we show those of the EC1 and EC2 clans from the Dominica dataset [13]. Only the most numerous coda types are shown: the rest of the vocal repertoires consists of more coda types with residual frequencies. C The subcoda structure can be modeled by considering rhythmic variations within codas of the same coda type. To do so, codas are represented as sequences of discrete inter-click intervals (dICIs), by discretizing absolute ICIs into discrete bins, which can then be considered akin to symbols (e.g. A, B, C, …), providing a tokenization of codas. Different instances of a single coda type can correspond to slightly different dICI sequences. The resulting dICI sequences are modeled using variable-length Markov chains, which can be represented as subcoda trees. These trees are built starting from the root node. Each node represents a dICI, and a dICI is a child of another in the tree if its presence in the affects the output probability of the parent, that is the probability of the parent depends on whether the child node has been observed previously in the sequence; see the Methods for a thorough explanation. These trees can be built for an individual speaker or for a group of speakers (here represented by colors), and capture the statistical properties of rhythm variations within codas and in the transitions between those. In other words, the tree captures a vocal style—how they say what they say. The vocal styles of different groups of sperm whales can be quantitatively compared by calculating a distance between their subcoda trees, which is a statistical distance between the dICI output distributions.

Removed in the revised version: Sperm whales live in multi-tiered societies and have a complex vocal communication system [14]. They communicate through rhythmic patterns (codas) of short broadband sounds (clicks), which have traditionally been classified into a set of coda types. This classification is based on the total number of clicks and their rhythm and tempo extrapolated from the time interval between consecutive clicks: the inter-click interval (ICI) [15, 16] (Fig. 1A). For example, the 4-regular (4R2) type refers to a pattern of four evenly spaced clicks, whereas the 1+3 type refers to two clicks separated by a longer pause followed by two clicks in quick succession. Coda types are thus standardized rhythmic patterns, but individual vocalizations of a given coda type exhibit micro-variations around that pattern. In the case of sperm whales ICIs can be used instead of inter-onset intervals (IOIs) due to the very low variance in duration of the individual clicks [17, 18].

Earlier version: The set of vocalized coda types (coda usage) combined with how frequently each is vocalized (coda frequency) makes up a vocal repertoire (Fig. 1B). For example, the 4R2 coda is used by many sperm whales, but other coda types are more specific in their usage or frequency to certain groups of sperm whales. While there is evidence of individual variation in vocal repertoires [19, 20, 21], sperm whales belonging to the same social unit—a stable, matrilineally-based group of whales, share a common vocal repertoire that is stable across years [20, 21, 22]. Social units that share substantial parts of their repertoire are said to be part of the same vocal clan [23, 24]. There is clear social segregation between members of different clans, even when living in sympatry, and thus clans mark a higher level of social organization, which appears to be defined on the basis of cultural vocal markers [23, 24, 12] (see Table 1 for a summary of the key concepts).

Revised version: The set of vocalized coda types (coda usage) combined with how frequently each is vocalized (coda frequency) makes up a vocal repertoire (Fig. 1B). For example, the 4R2 coda is used by many sperm whales, but other coda types are more specific in their usage or frequency to certain groups of sperm whales. While there is evidence of individual variation in vocal repertoires [22, 23, 24], sperm whales belonging to the same social unit—a stable, matrilineally-based group of whales, share a common vocal repertoire that is stable across years [23, 24, 25]. Social units that share substantial parts of their repertoire are said to be part of the same vocal clan [26, 27]. There is clear social segregation between members of different clans, even when living in sympatry, and thus clans mark a higher level of social organization, which appears to be defined on the basis of cultural vocal markers [26, 27, 12] (see Table 1 for a summary of the key concepts).

The clan specific and frequent usage of certain coda types, termed identity codas [16, 12], align with the expectations for symbolic markers of group membership [10]. Furthermore, quantitative evidence that sperm whales themselves use identity codas as such markers has recently emerged: the more two clans overlap in geographic space, the more different their identity coda usage is [12]. This is consistent with computational models [10] of the evolution of symbolic marking, which predict that differences between cultural norms will be starkest when inter-group interactions are more common (e.g., in boundary or overlap regions).

All remaining coda types have been referred to as non-identity (non-ID) codas and constitute a very large fraction of sperm whales’ total number of coda utterances. In fact, the total number of emitted non-ID codas accounts for more than 6 out of 10 codas (see SM Section 1.1 for the counts per clan and per coda type). This begets the question: if ID codas are used as clan identity signals, what can be said about the remaining 65% of codas?

Here, we introduce a novel descriptive framework that focuses on the subcoda structure, that is, the rhythmic micro-variations of intervals between clicks within codas (Fig. 1C). This framework, formally encoded in what we call a “subcoda tree”, captures how codas are uttered: a vocal style. We find that variations in this vocal style, even for a single coda type, identify an individual’s social unit and clan, effectively fingerprinting vocal repertoires. With this, we add a new dimension with respect to previous approaches based on which codas are said-vocal repertoires. Thus, we propose a new concept of vocal identity of sperm whales that comprises both vocal style and vocal repertoire.

Earlier version: By applying our modeling framework to acoustic data from the Atlantic and Pacific Oceans, we obtain two main results. First, we partition sperm whale populations into vocal style-defined clans, which we find to recapitulate the previously defined vocal repertoire clans. This confirms our method’s validity as it reiterates previous studies’ results. Second, and crucially, we find that the vocal style of non-ID codas is more similar for more sympatric clans, i.e. clans whose territory overlaps more spatially. In contrast, we do not find an effect of sympatry on the similarity of vocal styles when studying only ID codas. This suggests that geographic overlap induces vocal styles to become more similar between clans, without jeopardizing each clan’s acoustic identity signals. Our results strengthen previous results on the use of ID codas as symbolic markers, while supporting cultural transmission and social learning of vocalizations among whales of different clans, as predicted by theoretical models [25].

Revised version: By applying our modeling framework to acoustic data from the Atlantic and Pacific Oceans, we obtain two main results. First, we partition sperm whale populations into vocal style-defined clans, which we find to recapitulate the previously defined vocal repertoire clans. This confirms our method’s validity as it reiterates previous studies’ results. Second, and crucially, we find that the vocal style of non-ID codas is more similar for more sympatric clans, i.e. clans whose territory overlaps more spatially. In contrast, we do not find an effect of sympatry on the similarity of vocal styles when studying only ID codas. This suggests that geographic overlap induces vocal styles to become more similar between clans, without jeopardizing each clan’s acoustic identity signals. Our results strengthen previous results on the use of ID codas as symbolic markers, while supporting cultural transmission and social learning of vocalizations among whales of different clans, as predicted by theoretical models [28].

Results

Subcoda structure captures variability in sperm whale communication

We model the internal structure of codas, in terms of rhythmic variations at the level of clicks, by using variable length Markov chains (VLMCs). Our analytical pipeline is illustrated in Fig. 1C. We build each VLMC in two main steps. We first convert codas, naturally represented as sequences of continuous, absolute, inter-click intervals (ICIs), to sequences of discrete ICIs (dICIs), by discretizing time into bins. In this way, each dICI represents a narrow range of possible ICI values. The bins have a fixed width (or resolution) δt and thus implicitly correspond to the temporal resolution of our representation (see Methods for details on the optimal choice of δt). Note that although ICIs have units of time (seconds), dICIs are (unit-less) symbols (e.g. A, B, C, etc.), representing multiples of δt (and so the smaller δt, the more the symbols). For example, the shortest ICIs will be mapped to the symbol A whereas longer ones will be mapped to symbols further down the alphabet. Hence, each coda (a sequence of ICIs) is mapped to a sequence of discrete symbols (a sequence of dICIs). The second part of the pipeline focuses on modeling the internal structure of codas in terms of dICI sequences. Essentially, we want to estimate transition probabilities from a dICI sub-sequence to the next dICI. A standard way would be to describe this using k-order Markov chain models, which model the next state (dICI) probability given a past sequence of k-states (dICIs). However, it is possible that specific dICIs depend on sequences of varying length. Despite we do not have direct evidence of unitary blocks in sperm whale communication, on can imagine this effect similarly to what happens with words (e.g., a word beginning with “re” can continue in more ways than one starting with “zy”). To account for the possibility that states (words, dICIs) might depend on sequences of varying length, while also retaining only the most compressed statistical representation of how codas are structured in terms of dICIs, we employ VLMCs.

VLMCs are generalizations of standard (fixed-memory) Markov chains that allow sub-sequences of dICIs of variable lengths. Longer sequences are kept only if they are significantly more informative in predicting the next dICI than random chance, yielding an optimally compressed representation (see Methods for details on model fitting and selection, including the optimal choice of δt). Furthermore, VLMCs naturally have a tree structure (see Fig. 1C), because of the natural order between sequences and their sub-sequences. In particular, each node represents a sub-sequence of dICIs, and is equipped with a probability distribution of transitions to the next dICIs. The origin node corresponds to the empty sequence, leaf nodes correspond to the longest sequences, and all nodes forming the branch in between correspond to the sub-sequences of that leaf node. Thus, we call VLMCs fitted to coda ICI data subcoda trees.

Note that dICI sequences encode rhythmic variations within codas. Indeed, a coda type is a standard rhythmic pattern that can be realized with variations in its ICIs and thus in its dICIs too. For example, the 4R2 coda type can be vocalized as BCC but also as CBB (in a representation with, say, 26 symbols). In that sense, subcoda trees, through the dICIs sequences that they contain and their transition probabilities, capture information about a vocal style. Four more features of subcoda trees are noteworthy: (i) because the method’s input is a set of codas, we can build subcoda trees for repertoires corresponding to different social scales, from individual sperm whales, to social units, all the way up to vocal clans; (ii) the difference between different subcoda trees can be measured using a probabilistic distance (see Methods), which we can use to compare subcoda trees across sperm whale clans; (iii) certain features of the vocal style can be quantified via metrics on the subcoda tree: for example we can define a complexity of the vocal style measured by an entropy on the tree; and (iv) subcoda trees can also be used as generative models, to create new synthetic codas in the form of dICI sequences to train downstream machine learning models.

Added in the revised version: Vocal style recovers vocal clan structure

Added in the revised version: The information about vocal style contained in subcoda trees is sufficient to recover the social structure of sperm whales (social units and clans). We show this in two ways. First, we analyze a dataset from sperm whales in Dominica (Dominica dataset) [27]. This dataset has rich annotations (coda type annotations, identity of recorded whales, social relations of recorded whales) which makes it particularly useful for validation. Specifically, the sperm whales in the Dominica dataset are divided into well known social clans, each composed of several social units,each with its own specific vocal repertoires, and thus can be defined as two different vocal clans. For each social unit in this dataset, we aggregate the individual whales’ coda samples and build a subcoda tree. Computing the distance between these trees (see Methods), we find that the distances between social units within the same clan are significantly smaller than between clans (Fig. 2A). We also find that an agglomerative clustering (average linkage, see Methods for details) on the distance between the subcoda trees correctly clusters social units into their respective clans (Fig. 2B). Without a priori knowledge of the clan memberships, we used vocal style to recover the existing classification of social units into two clans, which was previously done based on similarity between vocal repertoires (i.e., coda types and usage) [27].

Vocal style recovers social structure of vocal clans in Dominica sperm whales.

Earlier version: A We show the similarity of vocal style, measured as subcoda tree-distance, among social units within a vocal clan (within, darker color shade) and between two clans (between, lighter color shade). We used the manual clan assignments from [24] as ground truth. Vocal style is more similar within clans than between clans. B We show the hierarchical clustering of social unit subcoda trees. Each leaf corresponds to a social unit, and the colors below show their known clan assignments. The clustering recovers the two-clan structure observed in past work [24].

Revised version: A We show the similarity of vocal style, measured as subcoda tree-distance, among social units within a vocal clan (within, darker color shade) and between two clans (between, lighter color shade). We used the manual clan assignments from [27] as ground truth. Vocal style is more similar within clans than between clans. B We show the hierarchical clustering of social unit subcoda trees. Each leaf corresponds to a social unit, and the colors below show their known clan assignments. The clustering recovers the two-clan structure observed in past work [27].

Removed in the revised version: Vocal style recovers vocal clan structure

Removed in the revised version: The information about vocal style contained in subcoda trees is sufficient to recover the social structure of sperm whales (social units and clans). We show this in two ways. First, we analyze a dataset from sperm whales in Dominica (Dominica dataset) [24]. This dataset has rich annotations (coda type annotations, identity of recorded whales, social relations of recorded whales) which makes it particularly useful for validation. Specifically, the sperm whales in the Dominica dataset are divided into well known social clans, each composed of several social units,each with its own specific vocal repertoires, and thus can be defined as two different vocal clans. For each social unit in this dataset, we aggregate the individual whales’ coda samples and build a subcoda tree. Computing the distance between these trees (see Methods), we find that the distances between social units within the same clan are significantly smaller than between clans (Fig. 2A). We also find that an agglomerative clustering (average linkage, see Methods for details) on the distance between the subcoda trees correctly clusters social units into their respective clans (Fig. 2B). Without a priori knowledge of the clan memberships, we used vocal style to recover the existing classification of social units into two clans, which was previously done based on similarity between vocal repertoires (i.e., coda types and usage) [24].

Second, we find that the subcoda structure of synthetic codas, generated from subcoda trees fitted on real data, closely reproduces that of real codas. To do this, we first train a simple classifier to assign codas to one of the two vocal clans, based on coda type. Variations of the same classifier, trained on the same real data, have been shown to discriminate between individual whales, social units, and clans with high accuracy [13]. We train the classifier on real codas, and then test it on both real and synthetic ones. The synthetic codas were generated using the subcoda tree of each clan, with a number of codas similar to that of the original dataset for a fair comparison (see Methods for details). We find that synthetic codas are correctly classified into their clans with an accuracy close (~ 85%) to that obtained on the real data (~ 90%, see Supplementary Materials Section 4).

Motivated by these results, we extend our analysis to a much larger dataset from the Pacific Ocean (Pacific dataset) [12]. This dataset is more sparsely annotated because of the breadth of its spatial coverage. We restricted our analyses to a well-sampled subset (n = 57 coda samples) of the full Pacific dataset (see Methods for details). Coda samples are only labeled by the spatial position at which they were recorded, but no information is available about the identity of the vocalizing sperm whales (see Methods for details). In fact, each repertoire likely contains codas from multiple individuals of a single clan. It has recently been shown that these coda samples can be divided into seven vocal clans based on their coda usage [12]. We use those clans as a benchmark for the following analysis.

Added in the revised version: Since there is no social unit-level information for this dataset, we fit a subcoda tree for each repertoire (i.e., all of the codas recorded on a single day in a single region). Trees are significantly more similar for coda samples belonging to the same vocal clan than for those belonging to different vocal clans (Fig. 3A). We also find that clustering coda samples based on vocal style returns a dendrogram that closely matches the one obtained from coda usage in [12] (Fig. 3B). The ma jor exception we find is the Short clan (red), named because member whales produce short codas with very few clicks, for which anomalous results were previously reported as well [12]. In our case, this is due to the Short clan being less well localized in the space of trees, while the other clans have well-defined centroids (see Supplementary Materials Fig. 9 for a low-dimensional representation subcoda tree metric space).

Vocal style recovers social structure of vocal clans in Pacific Ocean sperm whales.

A We show the similarity of vocal style, measured as subcoda tree-distance, between coda samples within a vocal clan (within, darker color shade) and between a clan and all others (between, lighter color shade). We used the vocal clans identified in [12]. Vocal style is more similar within a clan than between clans. B We show the hierarchical clustering of subcoda trees. Each leaf corresponds to a coda sample, and the colors below show their vocal clan assignments (based on coda usage) from [12]. We find generally good overlap between the groups obtained from clustering vocal style and those from vocal repertoire, with the exception of the Short clan (red) that is somewhat mixed with the Palindrome (orange) and Rapid Increasing (yellow) clans.

Removed in the revised version: Since there is no social unit-level information for this dataset, we fit a subcoda tree for each repertoire (i.e., all of the codas recorded on a single day in a single region). Trees are significantly more similar for coda samples belonging to the same vocal clan than for those belonging to different vocal clans (Fig. 3A). We also find that clustering coda samples based on vocal style returns a dendrogram that closely matches the one obtained from coda usage in [12] (Fig. 3B). The major exception we find is the Short clan (red), named because member whales produce short codas with very few clicks, for which anomalous results were previously reported as well [12]. In our case, this is due to the Short clan being less well localized in the space of trees, while the other clans have well-defined centroids (see Supplementary Materials Fig. 9 for a low-dimensional representation subcoda tree metric space).

Earlier version: Therefore, we find that sperm whale vocal clans in the Atlantic Ocean (Caribbean Sea) and Pacific Ocean can be identified by a vocal identity that encompasses both clan-specific vocal repertoire [24, 27, 23, 12] and vocal style as defined in this work.

Revised version: Therefore, we find that sperm whale vocal clans in the Atlantic Ocean (Caribbean Sea) and Pacific Ocean can be identified by a vocal identity that encompasses both clan-specific vocal repertoire [27, 30, 26, 12] and vocal style as defined in this work.

Clan sympatry impacts vocal style of non-ID codas only

While interesting, the fact that both vocal repertoires and vocal styles discriminate between clans might imply that considering both could be redundant for vocal identity. However, we find that this is not the case when we consider the functional role of ID versus non-ID codas.

More precisely, different clans can share significant portions of their total range, overlapping across large swaths of ocean. Such sympatric clans exhibit a decreasing similarity of their ID coda usage with increasing clan overlap [12]. This means that the more two clans overlap in space, the more dissimilar their vocal repertoires are in terms of ID coda types and their usage frequency. This is consistent with the idea that ID codas are used as symbolic markers to delineate cultural boundaries between social groups [12, 10]. In contrast, non-ID coda usage do not show any relationship to clan overlap.

Added in the revised version: We find the exact opposite effect when considering vocal style. The similarity in vocal style for ID codas across clans does not depend on the level of clan overlap (Fig. 4a). In contrast, the similarity in vocal style for non-ID codas displays a clear and significant increase (i.e., decreasing subcoda tree-distance) as clan spatial overlap increases (Fig. 4b). In the Supplementary Materials (see Section 2.4.2), we show that the same results hold at the single coda type level, in addition to the whole clan level, along with an analysis of the confidence intervals. These results imply that the internal structure of codas is more similar for groups that likely spend more time in the same space, which can be thought to be analogous to accents aligning in human populations that share the same territory [31, 32]. This also highlights the complementarity of vocal repertoire and style: the trends are different precisely because the two concepts describe different aspects of whale vocalization.

Clan overlap influences non-ID coda vocal styles

Comparing the similarities of different VLMC models fit for each Pacific Ocean clan for both ID and non-ID coda samples. The y-axis represents the measured distance between the subcoda trees, and the x-axis shows the geographical clan overlap (as calculated in [12]; see “Measuring clan overlap” in our Methods for details). Each point represents a pairwise comparison between two clans. The effect of overlap on ID coda vocal style similarity is minimal and non-significant while the opposite is true for non-ID codas: overlapping clans produce non-ID codas with a more similar vocal style. The VLMC distances are also typically much greater for ID codas than for non-ID codas. Note that while these results are visually opposite to those reported in Hersh et al. [12], they support the same final conclusions (see “Identity and non-identity codas show different trends” for details).

Removed in the revised version: We find the exact opposite effect when considering vocal style. The similarity in vocal style for ID codas across clans does not depend on the level of clan overlap (Fig. 4a). In contrast, the similarity in vocal style for non-ID codas displays a clear and significant increase (i.e., decreasing subcoda tree-distance) as clan spatial overlap increases (Fig. 4b). In the Supplementary Materials (see Section 2.4.2), we show that the same results hold at the single coda type level, in addition to the whole clan level, along with an analysis of the confidence intervals. These results imply that the internal structure of codas is more similar for groups that likely spend more time in the same space, which can be thought to be analogous to accents aligning in human populations that share the same territory [28, 29]. This also highlights the complementarity of vocal repertoire and style: the trends are different precisely because the two concepts describe different aspects of whale vocalization.

Discussion

Earlier version: We have presented a general method for modeling animal communication systems and their complexity based on VLMCs. In the context of sperm whales, this new method allows the extraction of subcoda trees, which succinctly describe the internal temporal structure of codas. Previous work on the structure of sperm whale communication has largely focused on supra-level coda analyses: for example, by classifying codas into types, quantifying how often different types are used, and distinguishing between individual whales, social units, or clans based on those counts [21, 30]. Here, we adopted a more fine-scale approach by investigating potential structure within codas. To do so, we used VLMCs to model the transition probability of observing a specific ICI given the previous ones. A VLMC, or here a subcoda tree, encodes all those probabilities but only for dICI sequences that are informative—other sequences are automatically discarded. As such, a subcoda tree is a statistically validated representation of the internal memory structure of codas at the level of sequences of clicks. It contains information about important rhythmic variations and transitions between them: a vocal style.

Revised version: We have presented a general method for modeling animal communication systems and their complexity based on VLMCs. In the context of sperm whales, this new method allows the extraction of subcoda trees, which succinctly describe the internal temporal structure of codas. Previous work on the structure of sperm whale communication has largely focused on supra-level coda analyses: for example, by classifying codas into types, quantifying how often different types are used, and distinguishing between individual whales, social units, or clans based on those counts [24, 33]. Here, we adopted a more fine-scale approach by investigating potential structure within codas. To do so, we used VLMCs to model the transition probability of observing a specific ICI given the previous ones. A VLMC, or here a subcoda tree, encodes all those probabilities but only for dICI sequences that are informative—other sequences are automatically discarded. As such, a subcoda tree is a statistically validated representation of the internal memory structure of codas at the level of sequences of clicks. It contains information about important rhythmic variations and transitions between them: a vocal style.

Using such representations, we propose a novel concept of vocal identity for sperm whales composed of vocal repertoire (what they say) and vocal style (how they say it), the latter being captured by our framework. We find that: (i) vocal styles vary between social units and clans, and can be used to distinguish them; (ii) the similarity of clan vocal styles for non-ID codas increases with increasing spatial overlap, while no change occurs for ID codas; and (iii) social learning across symbolic cultural boundaries most parsimoniously explains the observed trends.

Vocal style recovers hierarchical social structure

Earlier version: Using the Dominica dataset, sperm whales had previously been divided into two vocal clans, based on their vocal repertoires and observed social interactions [24]. In our study, comparing the vocal styles of those same whales led to the same assignment of social units to two vocal clans. Similarly, for the Pacific dataset, clustering based on vocal styles yielded clans that were in good agreement with those previously defined based on vocal repertoires [12] (Supplementary Materials for an extended comparison). The difference between the two partitions was mainly due to the Short clan, which was more spread out in subcoda tree space than the other clans, causing overlap with other clans that showed less variability. This variability could be linked to the fact that Short clan whales typically make codas with very few (e.g., three or four) clicks, leading to subcoda trees with very few nodes. In Ref. [12], the authors observe a similar lack of uniformity in coda usage of the Short clan.

Revised version: Using the Dominica dataset, sperm whales had previously been divided into two vocal clans, based on their vocal repertoires and observed social interactions [27]. In our study, comparing the vocal styles of those same whales led to the same assignment of social units to two vocal clans. Similarly, for the Pacific dataset, clustering based on vocal styles yielded clans that were in good agreement with those previously defined based on vocal repertoires [12] (Supplementary Materials for an extended comparison). The difference between the two partitions was mainly due to the Short clan, which was more spread out in subcoda tree space than the other clans, causing overlap with other clans that showed less variability. This variability could be linked to the fact that Short clan whales typically make codas with very few (e.g., three or four) clicks, leading to subcoda trees with very few nodes. In Ref. [12], the authors observe a similar lack of uniformity in coda usage of the Short clan.

Identity and non-identity codas show different trends

For ID codas, we show that the similarity between clan vocal styles is not affected by spatial overlap, while it has recently been shown that the similarity between clan vocal repertoires decreases with overlap (measured on the two clans’ joint set of ID codas) [12]. This means that spatial overlap does not affect how whales produce ID codas (in terms of their fine-scale rhythmic structure; our results) but does affect how often they produce them [12]. In contrast, for non-ID codas, we show that the similarity between vocal styles increases with spatial overlap between two clans, while no change was observed for vocal repertoires in previous work on the same dataset. In other words, increasing spatial overlap is correlated with more similar fine-scale rhythmic structure of non-ID codas produced by whales from different clans (our results), but does not affect how often non-ID codas are produced. Our study thus supports and nuances the results of Hersh et al. [12]. We provide further support for selection acting to produce unambiguous, recognizable identity signals in the ID codas. However, ID codas only account for 35% of the total vocalizations; the remaining 65% of codas have traditionally been lumped into a catch-all category (i.e., non-ID codas) and their function remains enigmatic (these numbers are an average over the Pacific clans, and go up to 93% for non-ID codas when counting number of coda types instead of number of codas emitted, see SM 1.1 for details). We could still discriminate among clans using non-ID coda vocal style; however, the increased similarity of non-ID coda vocal styles between clans with greater spatial overlap, as demonstrated here, suggests that non-ID codas are likely vocal cues and not identity signals like the ID codas. Accordingly, vocal repertoire and vocal style capture different and complementary information on sperm whale communication, and should be considered in tandem in future studies.

Evidence for social learning across cultural boundaries

There are several potential mechanisms driving the similarity in non-ID coda vocal styles—but not ID coda vocal styles-across spatially overlapped clans: environmental variation, genetics, and/or social learning.

Earlier version: Local adaptation to specific ecological conditions can lead to geographic variation in acoustic signals [31]. If environmental pressures alone were responsible for the trends we observe in sperm whales, this would imply that (i) more spatially overlapped clans experience more similar environments, (ii) non-ID coda vocal style is impacted by or dependent on environmental parameters, and (iii) ID coda vocal style is not impacted by/dependent on environmental parameters. Although the first point is somewhat intuitive, to date there is no evidence that coda production systematically varies with environment. In fact, clans are recognizable across ocean basins, making local adaptation an unlikely driver of the observed trend in non-ID coda vocal style.

Revised version: Local adaptation to specific ecological conditions can lead to geographic variation in acoustic signals [34]. If environmental pressures alone were responsible for the trends we observe in sperm whales, this would imply that (i) more spatially overlapped clans experience more similar environments, (ii) non-ID coda vocal style is impacted by or dependent on environmental parameters, and (iii) ID coda vocal style is not impacted by/dependent on environmental parameters. Although the first point is somewhat intuitive, to date there is no evidence that coda production systematically varies with environment. In fact, clans are recognizable across ocean basins, making local adaptation an unlikely driver of the observed trend in non-ID coda vocal style.

Earlier version: If genetic relatedness were responsible, this would imply that (i) more spatially overlapped clans are more genetically related, (ii) non-ID coda vocal styles are genetically inherited, and (iii) ID coda vocal styles are not genetically inherited. If all three requirements were met, then the observed similarity in non-ID coda vocal styles for more spatially overlapped clans could be due to genetic determination under a general isolation by distance structure. However, research to date suggests this scenario is unlikely. Rendell et al. [32] found little evidence to support genetics as an explanation of differences in vocal dialects among clans in the Pacific Ocean. Furthermore, Alexander et al. [33] found that regional genetic differentiation in the Pacific Ocean is very low: while social group is important for explaining both mitochondrial and nuclear DNA variance, geographic region is not. This contrasts with results from the Indian Ocean, where region was the strongest predictor of mitochondrial DNA variance. Given that gene flow in sperm whales is largely male-mediated and that mitochondrial DNA haplotypes are broadly shared across the Pacific Ocean, it is unlikely that coda dialects are genetically determined [33, 34]. Agent-based models grounded in empirical data from Pacific Ocean sperm whales further support coda usage as socially learned, not genetically inherited [25]. To fully rule out a genetic explanation for our results, the analyses in [32] could be replicated for ID coda usage and non-ID coda usage separately. This would shed light on whether certain coda types are genetically inherited vs. socially learned, as has been suggested for some humpback whale (Megaptera novaeangliae) vocalizations [35].

Revised version: If genetic relatedness were responsible, this would imply that (i) more spatially overlapped clans are more genetically related, (ii) non-ID coda vocal styles are genetically inherited, and (iii) ID coda vocal styles are not genetically inherited. If all three requirements were met, then the observed similarity in non-ID coda vocal styles for more spatially overlapped clans could be due to genetic determination under a general isolation by distance structure. However, research to date suggests this scenario is unlikely. Rendell et al. [35] found little evidence to support genetics as an explanation of differences in vocal dialects among clans in the Pacific Ocean. Furthermore, Alexander et al. [36] found that regional genetic differentiation in the Pacific Ocean is very low: while social group is important for explaining both mitochondrial and nuclear DNA variance, geographic region is not. This contrasts with results from the Indian Ocean, where region was the strongest predictor of mitochondrial DNA variance. Given that gene flow in sperm whales is largely male-mediated and that mitochondrial DNA haplotypes are broadly shared across the Pacific Ocean, it is unlikely that coda dialects are genetically determined [36, 37]. Agent-based models grounded in empirical data from Pacific Ocean sperm whales further support coda usage as socially learned, not genetically inherited [28]. To fully rule out a genetic explanation for our results, the analyses in [35] could be replicated for ID coda usage and non-ID coda usage separately. This would shed light on whether certain coda types are genetically inherited vs. socially learned, as has been suggested for some humpback whale (Megaptera novaeangliae) vocalizations [38].

Earlier version: The most parsimonious explanation for the observed similarity of non-ID coda vocal styles of clans with increasing spatial overlap is social learning across clan boundaries. This is remarkable, given that sperm whales belonging to different clans have rarely been observed physically interacting at sea [4]. However, that does not preclude the possibility that they are within acoustic range of each other [36] and that cross-cultural social learning opportunities arise. This explanation is compatible with (and bolsters) past work suggesting that ID and non-ID codas function differently in sperm whale communication, and further suggests that they experience different evolutionary pressures [25]. Whether social learning has facilitated stochastic (i.e., cultural drift) or deterministic (i.e., cultural selection) processes is more difficult to determine, and it is unclear whether the observed non-ID coda vocal style alignment has been neutral or adaptive [37, 31]. Importantly, these findings suggest that vocal learning in sperm whales may not be limited to vertical transmission from related adults to young kin, but that horizontal and/or oblique social learning from outside the natal social unit might also be occuring.

Revised version: The most parsimonious explanation for the observed similarity of non-ID coda vocal styles of clans with increasing spatial overlap is social learning across clan boundaries. This is remarkable, given that sperm whales belonging to different clans have rarely been observed physically interacting at sea [4]. However, that does not preclude the possibility that they are within acoustic range of each other [39] and that cross-cultural social learning opportunities arise. This explanation is compatible with (and bolsters) past work suggesting that ID and non-ID codas function differently in sperm whale communication, and further suggests that they experience different evolutionary pressures [28]. Whether social learning has facilitated stochastic (i.e., cultural drift) or deterministic (i.e., cultural selection) processes is more difficult to determine, and it is unclear whether the observed non-ID coda vocal style alignment has been neutral or adaptive [40, 34]. Importantly, these findings suggest that vocal learning in sperm whales may not be limited to vertical transmission from related adults to young kin, but that horizontal and/or oblique social learning from outside the natal social unit might also be occuring.

Earlier version: Vocal identity in sperm whales is thus consistent with both cultural selection on ID codas to maintain discrete signals for vocal recognition in sympatry, and social learning between clans leading to a vocal style more similar to that of other whales with which they are in acoustic contact more frequently. This highlights a more complex system of transmission in which clan identity is maintained through selection, while gradual change over time may occur within and across clans for vocalizations which do not function in social recognition and thus may create similar vocal styles. Indeed, the observed alignment of non-ID vocal styles may reflect a combination of automatic imitation, and controlled imitation, guided by social and ecological factors, as discussed in studies of social learning mechanisms [38, 39]. In the context of our findings, the similarity in non-identity codas across clans with geographic overlap may indicate an automatic component of social learning. Whales sharing the same habitat are exposed to overlapping acoustic environments, which could facilitate the development of similar vocal styles described here. However, the persistence of distinct identity codas within clans may suggests a more directed form of social learning, shaped by social goals such as clan recognition and cohesion.

Revised version: Vocal identity in sperm whales is thus consistent with both cultural selection on ID codas to maintain discrete signals for vocal recognition in sympatry, and social learning between clans leading to a vocal style more similar to that of other whales with which they are in acoustic contact more frequently. This highlights a more complex system of transmission in which clan identity is maintained through selection, while gradual change over time may occur within and across clans for vocalizations which do not function in social recognition and thus may create similar vocal styles. Indeed, the observed alignment of non-ID vocal styles may reflect a combination of automatic imitation, and controlled imitation, guided by social and ecological factors, as discussed in studies of social learning mechanisms [41, 42]. In the context of our findings, the similarity in non-identity codas across clans with geographic overlap may indicate an automatic component of social learning. Whales sharing the same habitat are exposed to overlapping acoustic environments, which could facilitate the development of similar vocal styles described here. However, the persistence of distinct identity codas within clans may suggests a more directed form of social learning, shaped by social goals such as clan recognition and cohesion.

Future directions

Earlier version: Our results can be expanded in multiple ways in future work. The first, and the simplest conceptually, would be to conduct the present analysis on a larger dataset. More codas would improve the quality of the statistical analyses and ensure that all codas are represented in realistic proportions for each clan. Moreover, longitudinal datasets might provide direct evidence to discriminate between the social learning hypothesis and competing ones (e.g. drift in vocal style). Similarly, confirmations could emerge from large scale genetic datasets addressing the issues of phylogenetic relatedness (or lack thereof) in clans that are closer in vocal style distance. Such datasets do not exist at present, but efforts towards automated and semi-automated collection techniques are underway (e.g. Project CETI [40]). Second, from a methodological perspective, we could add spectral information (in terms of acoustic frequencies) to the temporal information currently used. Although sperm whale acoustic communication seems mostly based on rhythm, spectral features of individual clicks may convey additional information. This possibility could be incorporated into our method by labeling the dICIs according to the frequency content of the associated click (or by extending the available “alphabet” for the VLMC). Third, it would be interesting to investigate in more detail the function of non-ID codas. Indeed, even though ID codas were only recently formally named for the first time, they have been the primary focus of sperm whale coda research for decades. As previously mentioned, non-ID codas are a catch-all category for anything that is not an ID coda, but that does not mean that all non-ID codas function in the same way. To start to unveil their function, we need to consider the context (behavioral, environmental, etc.) in which different non-ID codas are produced [41]. The pattern we documented may or may not apply to all non-ID codas, but it is at least strong enough that we detect the relationship with clan spatial overlap when collectively considering all non-ID codas. Finally, the methods we developed here could in principle be used to capture the structural complexity of any communication system, including those coming from other animal species, provided they can be encoded in sequences of tokens (e.g. calls, signs, etc.). This would also allow to compare the complexity of communication systems across species, possibly linking it to other ecological parameters of the environments inhabited by such species.

Revised version: Our results can be expanded in multiple ways in future work. The first, and the simplest conceptually, would be to conduct the present analysis on a larger dataset. More codas would improve the quality of the statistical analyses and ensure that all codas are represented in realistic proportions for each clan. Moreover, longitudinal datasets might provide direct evidence to discriminate between the social learning hypothesis and competing ones (e.g. drift in vocal style). Similarly, confirmations could emerge from large scale genetic datasets addressing the issues of phylogenetic relatedness (or lack thereof) in clans that are closer in vocal style distance. Such datasets do not exist at present, but efforts towards automated and semi-automated collection techniques are underway (e.g. Project CETI [43]). Second, from a methodological perspective, we could add spectral information (in terms of acoustic frequencies) to the temporal information currently used. Although sperm whale acoustic communication seems mostly based on rhythm, spectral features of individual clicks may convey additional information. This possibility could be incorporated into our method by labeling the dICIs according to the frequency content of the associated click (or by extending the available “alphabet” for the VLMC). Third, it would be interesting to investigate in more detail the function of non-ID codas. Indeed, even though ID codas were only recently formally named for the first time, they have been the primary focus of sperm whale coda research for decades. As previously mentioned, non-ID codas are a catch-all category for anything that is not an ID coda, but that does not mean that all non-ID codas function in the same way. To start to unveil their function, we need to consider the context (behavioral, environmental, etc.) in which different non-ID codas are produced [44]. The pattern we documented may or may not apply to all non-ID codas, but it is at least strong enough that we detect the relationship with clan spatial overlap when collectively considering all non-ID codas. Finally, the methods we developed here could in principle be used to capture the structural complexity of any communication system, including those coming from other animal species, provided they can be encoded in sequences of tokens (e.g. calls, signs, etc.). This would also allow to compare the complexity of communication systems across species, possibly linking it to other ecological parameters of the environments inhabited by such species.

Methods

Acoustic data

We analyzed two datasets in the present study. The Dominica dataset contains 8719 annotated codas recorded in the Atlantic Ocean off the island of Dominica between 2005 and 2019. The codas come from 12 social units grouped into two vocal clans (EC1 and EC2). The Pacific dataset was collected between 1978 and 2017 at 23 locations in the Pacific Ocean (the recording methods are available in the supplementary materials of [12]). The codas were divided into coda samples according to their recording day and each repertoire was assigned a single vocal clan inferred in [12]. When considering a clan-level analysis (Fig. 3) all coda samples were used to compute the subcoda trees (23555 codas). However, when analysing at a coda samples level (Fig. 4), we discarded coda samples with less than 200 codas with statistical inference in mind, resulting in a final count of 57 coda samples (17046 codas) for the Pacific.

Representation of sperm whale communication as discrete inter-click intervals

Earlier version: As a preliminary step, we discretized the (continuous) ICI values into bins of width St seconds. In other words, we represented the continuum of ICI values by a finite set of discrete ICIs (dICIs) based on the duration of the ICI. The bin width St controls the temporal resolution of the representation: a higher value of St implies a coarser representation with fewer dICIs. We also imposed an upper bound tmax: any ICI value greater than that was truncated to tmax. This ensured that the set of dICIs was finite. Note that although ICIs have units of time (seconds), dICIs are unitless (they represent time intervals). The resulting representation of ICIs as dICIs is a discrete random variable defined as Xδt,tmax=[min(ICI, tmax)δt], (1)

Revised version: As a preliminary step, we discretized the (continuous) ICI values into bins of width δt seconds. In other words, we represented the continuum of ICI values by a finite set of discrete ICIs (dICIs) based on the duration of the ICI. The bin width δt controls the temporal resolution of the representation: a higher value of δt implies a coarser representation with fewer dICIs. We also imposed an upper bound tmax: any ICI value greater than that was truncated to tmax. This ensured that the set of dICIs was finite. Note that although ICIs have units of time (seconds), dICIs are unitless (they represent time intervals). The resulting representation of ICIs as dICIs is a discrete random variable defined as

which takes values in the finite set 𝒳={0,1,…,⌊tmaxδt⌋} . We represented the sequences of ICIs by sequences of dICIs from that finite set. Note that any ICI value above tmax is mapped to the dICI ⌊tmaxδt⌋ and therefore represents the end of a coda. We set tmax = 1 (longer than any ICI) and δt = 0.05 throughout the analysis (see Supplementary Materials section 3.3.2 for justification of this choice and section 3.4.3 for an analysis on the influence of this parameter).

Variable length Markov chains

We then modeled these dICI sequences using variable length Markov chains (VLMCs). VLMCs provide the large memory advantage of higher-order Markov chains when needed, without the drawback of having too many unnecessary parameters in the model.

Fitting a VLMC is the process of deciding how much memory is necessary to model specific sequences. The criterion for making this decision is the following: longer sequences are discarded if their distribution of transition probabilities is similar to that of shorter subsequences. This process is often called context tree estimation and consists of two steps.

Earlier version: The first step is to consider WD the set of all sequences of maximum length D (which we set to 10) and to assign the following probability distribution qw to each sequence: qw=P(X|w), (2)

Revised version: The first step is to consider WD the set of all sequences of maximum length D (which we set to 10) and to assign the following probability distribution qw to each sequence:

that is, the probability of observing a state x ∈ 𝒳 given the sequence w.

Earlier version: The second step is to prune the sequences that do not add information. Take two sequences u, w ∈ WD, one being the suffix of the other w = σu. The information gained Hw by considering the longer sequence can be measured with a weighted Kullback-Leibler (KL) divergence DKL [42]. The longer memory sequence w is kept only if the information gain is greater than some threshold K [43, 44] ΔHw=N(w)DKL(qw||qu)>K (3)

Revised version: The second step is to prune the sequences that do not add information. Take two sequences u, w ∈ 𝒲D, one being the suffix of the other w = σu. The information gained Hw by considering the longer sequence can be measured with a weighted Kullback-Leibler (KL) divergence DKL [45]. The longer memory sequence w is kept only if the information gain is greater than some threshold K [46, 47]

where N (w) denotes the length of sequence w. Sequences that satisfy this condition are called contexts and sequences that do not are discarded. A VLMC can be defined as the set of these contexts w and their associated probability distribution qw (see Supplementary Materials section 3.1 for details).

A VLMC can be visualized as a tree by representing each context w by a node and setting the root node as the context of length zero. Contexts that are subsequences of each other are then part of the same branches, which end with the longest contexts.

Quantitative Comparison of VLMCs

Earlier version: If two VLMC models T1 and T2 are built over the same finite set of dICIs χ, there exists a map ϕ1 : WD → T1 that maps any sequence of elements of χ into the longest sequence present in T1, and similarly for T2. This map also induces a map between the probability distributions of T1 and T2. Given two distributions over the same set χ, we can measure how different they are with the KL divergence. Therefore, it is possible to define a dissimilarity between T1 and T2 by considering the average KL divergence over all sequences of T1 and their map ϕ1(T2) ⊆ T1 dKL(T1,T2)=1|T1|∑w∈T1DKL(qw||pϕ1(w)) (4)

Revised version: If two VLMC models T1 and T2 are built over the same finite set of dICIs 𝒳, there exists a map ϕ1 : 𝒲D → T1 that maps any sequence of elements of 𝒳 into the longest sequence present in T1, and similarly for T2. This map also induces a map between the probability distributions of T1 and T2. Given two distributions over the same set 𝒳, we can measure how different they are with the KL divergence. Therefore, it is possible to define a dissimilarity between T1 and T2 by considering the average KL divergence over all sequences of T1 and their map ϕ1(T2) ⊆ T1

Refer to the Supplmentary Materials section 3.4 for a more detailed explanation.

This results in a dissimilarity measure that captures not just the difference in emission distribution but also the structural differences of the associated context trees. When comparing the distribution of distances in Fig. 2A and Fig. 3A we performed a Kolmogorov–Smirnov test to test if the distances between social units/coda samples of the same clan and distances between social units/coda samples of different clans had come from the same distribution. For every pair, we can reject the hypothesis of the distances coming from the same distribution with 95% confidence.

Hierachical Clustering of VLMCs

Earlier version: The dendrograms in Fig. 2B and Fig. 3B were obtained by hierarchical clustering using average linkage on the set of subcoda trees (VLMCs). Since the distance is not symmetric, for agglomerative clustering we considered the symmetric distance: dKL(T1,T2)=max{dKL(T1,T2),dKL(T2,T1)}. (5)

Revised version: The dendrograms in Fig. 2B and Fig. 3B were obtained by hierarchical clustering using average linkage on the set of subcoda trees (VLMCs). Since the distance is not symmetric, for agglomerative clustering we considered the symmetric distance:

Measuring clan overlap

Earlier version: We used the clan spatial overlap values from [12]. Briefly, given two clans A and B, and the coda samples associated to them, the amount of geographical overlap of A in B was measured as the fraction of coda samples belonging to clan A that were recorded within 1000 kilometers of at least one repertoire of clan B. One thousand kilometers is the approximate annual home range span of sperm whales in the eastern tropical Pacific [45, 46].

Revised version: We used the clan spatial overlap values from [12]. Briefly, given two clans A and B, and the coda samples associated to them, the amount of geographical overlap of A in B was measured as the fraction of coda samples belonging to clan A that were recorded within 1000 kilometers of at least one repertoire of clan B. One thousand kilometers is the approximate annual home range span of sperm whales in the eastern tropical Pacific [48, 49].

Statistical Testing

On Fig. 2 and Fig. 3 we compare the distributions of distances between subcoda trees of coda samples/social units of the same clan (within) and of different clans (between). The purpose is to assess whether these distributions originate from the same underlying population. We employ both the Kolmogorov–Smirnov test and the T-test. The observed p-values were well below 0.01 for all clans. This allows us to confidently reject the hypothesis that there is no difference between the vocal style between different clans. For more information check the Supplementary Materials in section 3.4.2.

To assess the existence of a relationship between clan overlap and vocal style similarity, we applied an ordinary least squares linear regression model (OLS). We show the resulting p values of the OLS statistical test at the bottom left of each plot of Fig. 4 along with the observed r2 value. To assess whether there is true difference between the two cases, we also bootstrapped the linear regression calculation to obtain 95% confidence intervals for the slopes of the fits, resulting in both negative and positive values in the ID case, but only negative slope values for the non-ID case, thus confirming our interpretation.